---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## INSTEAD OF Triggers for DML Statements

```yaml
type: VideoExercise
key: 8601ed4b19
xp: 50
```

`@projector_key`
4ac6d7e2ecae5b05febe9a86521c1950

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 2d07ec999b
xp: 50
```

`@projector_key`
5ab499fb0d6cc075b848cd512a1f2fce
